## Before you start working on the exercise

- Use Python version 3.9 and above.
- It is highly recommended to create a virtual environment for this course. You can find resources on how to create a virtual environment on the ISIS page of the course.
- Make sure that no assertions fail or exceptions occur, otherwise points will be subtracted.
- Use all the variables given to a function unless explicitly stated otherwise. If you are not using a variable you are doing something wrong.
- Read the **whole** task description before starting with your solution.
- After you submit the notebook more tests will be run on your code. The fact that no assertions fail on your computer locally does not guarantee that you completed the exercise correctly.
- Please submit only the notebook file with its original name. If you do not submit an `ipynb` file you will fail the exercise.
- Edit only between YOUR CODE HERE and END YOUR CODE.
- Verify that no syntax errors are present in the file.
- Before uploading your submission, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel\Restart) and then run all cells (in the menubar, select Cell\Run All).

```python
import sys

if (3, 9) <= sys.version_info[:2] <= (3, 11):
    print("Correct Python version")
else:
    print(
        f"You are using a wrong version of Python: {'.'.join(map(str,sys.version_info[:3]))}"
    )
```

Output:
```
Correct Python version
```

<!-- grade_id: cell-e5951c7ec892701b, points: 0, grade: true, solution: false -->
```python
# This cell is for grading. DO NOT remove it

# Use unittest asserts 
import unittest

t = unittest.TestCase()
from pprint import pprint

# Helper assert function
def assert_percentage(val):
    t.assertGreaterEqual(val, 0.0, f"Percentage ({val}) cannot be < 0")
    t.assertLessEqual(val, 1.0, f"Percentage ({val}) cannot be > 1")


```

# Warm Ups

Before starting the homework sheet we recommend you finish these warm-up tasks. They won't bring any points but should help you to get familiar with Python code.

### Function and types (0 P)

Write a function using list comprehension that returns the types of list elements.

* The function should be called `types_of`
* The function expects a list as an input argument.
* The function should return a list with the types of the given list elements.
* Read the testing cell to understand how `types_of` is supposed to work.

<!-- grade_id: cell-8d004130189cffa5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def types_of(X : list):
    return [type(x) for x in X]
# YOUR CODE HERE

```

<!-- grade_id: cell-e082a23baca66faf, points: 0, grade: true, solution: false -->
```python
# Test type_of_two function
types = types_of([7, 0.7, "hello", True, (2, "s")])

assert isinstance(types, list)
t.assertEqual(types[0], int)
t.assertEqual(types[1], float)
t.assertEqual(types[2], str)
t.assertEqual(types[3], bool)
t.assertEqual(types[-1], tuple)
```

### Concatenation and enumerate (0 P)


Concatenate the strings from the array 'animals' into one string.

* Use: `counting +=` and string formatting (`f-strings`).
* Use `enumerate` to get the `i`th index.
* The result should look as follows: `'0: mouse | 1: rabbit | 2: cat | 3: dog |'`

***Note that this is not the most efficient way to concatenate strings in Python but part of this exercise is to showcase `for-loops`***

```python
animals = ["mouse", "rabbit", "cat", "dog"]
```

<!-- grade_id: cell-598f867f4109ac39, points: 0, grade: false, solution: true -->
```python
counting = "|"
for i, animal in enumerate(animals):
    # YOUR CODE HERE
    counting += f'{i}: {animal} |'
    # YOUR CODE HERE
    

print(counting)
```

Output:
```
|0: mouse |1: rabbit |2: cat |3: dog |
```

<!-- grade_id: cell-cbf9e2c30017f8d7, points: 0, grade: true, solution: false -->
```python
# Test of the enumeration loop
t.assertEqual(counting, "|0: mouse |1: rabbit |2: cat |3: dog |")
```

### String formating (0 P)

What does the following string formating result in?
* Write the result of the string formating into the variables result1, result2, result3.
* Example: `string0 = "This is a {} string.".format("test")`
* Example solution: `result0 = "This is a test string"`

```python
# first string
string1 = "The sky is {}. {} words in front of {} random words create {} random sentence.".format(
    "clear", "Random", "other", 1
)

# second string
a = "irony"
b = "anyone"
c = "room"

string2 = f"The {a} of the situation wasn't lost on {b} in the {c}."

# third string
string3 = f"{7*10} * {9/3} with three digits after the floating point looks like this: {70*3 :.3f}."

# fourth string
string4 = "   Hello World.   ".strip()
```

<!-- grade_id: cell-ab79852811b12118, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
result1 = 'The sky is clear. Random words in front of other random words create 1 random sentence.'
result2 = 'The irony of the situation wasn\'t lost on anyone in the room.'
result3 = '70 * 3.0 with three digits after the floating point looks like this: 210.000.'
result4 = 'Hello World.'
# YOUR CODE HERE

```

<!-- grade_id: cell-09a4a5a04b6d9098, points: 0, grade: true, solution: false -->
```python
# Test the string results
t.assertEqual(string1, result1)
t.assertEqual(string2, result2)
t.assertEqual(string3, result3)
t.assertEqual(string4, result4)
```

# Exercise Sheet 1: Python Basics

This first  exercise sheet tests the basic functionalities of the Python programming language in the context of a simple prediction task. We consider the problem of predicting health risk of subjects from personal data and habits. We first use for this task a decision tree.

![](tree.png)

Make sure that you have downloaded the `tree.png` file from ISIS. For this exercise sheet, you are required to use only pure Python, and to not import any module, including `Numpy`. Next week are going to implement the nearest neighbor part of this exercise sheet using `NumPy` 😉.

## Classifying a single instance (15 P)

* In this sheet we will represent patient info as a tuple.
* Implement the function `decision` that takes as input a tuple containing values for attributes (smoker,age,diet), and computes the output of the decision tree. Should return either `'less'` or `'more'`. No other outputs are valid.

<!-- grade_id: cell-b66d7278bc313c94, points: 0, grade: false, solution: true -->
```python
def decision(x):
    """
    This function implements the decision tree represented in the above image. As input the function
    receives a tuple with three values that represent some information about a patient.
    Args:
        x: Input tuple containing exactly three values.
        The first element represents a patient is a smoker this value will be 'yes'.
        All other values represent that the patient is not a smoker. The second
        element represents the age of a patient in years as an integer. The last
        element represents the diet of a patient. If a patient has a good diet
        this string will be 'good'. All other values represent that the patient
        has a poor diet.
    Returns:
        str: A string that has either the value 'more' or 'less'. No other return value
        is valid.

    """
    # YOUR CODE HERE
    if x[0] == 'yes':
        if x[1] < 29.5:
            return 'less'
        else:
            return 'more'
        
    else:
        if x[2] == 'good':
            return 'less'
        else:
            return 'more'
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-c31b80471db3132f, points: 15, grade: true, solution: false -->
```python
# Test decision function

# Test expected 'more'
x = ("yes", 31, "good")
output = decision(x)
print(f"decision({x}) --> {output}")
t.assertIsInstance(output, str)
t.assertEqual(output, "more")

# Test expected 'less'
x = ("yes", 29, "poor")
output = decision(x)
print(f"decision({x}) --> {output}")
t.assertIsInstance(output, str)
t.assertEqual(output, "less")

```

Output:
```
decision(('yes', 31, 'good')) --> more
decision(('yes', 29, 'poor')) --> less
```

<!-- grade_id: cell-a706e6de3303d3e7, points: 0, grade: true, solution: false -->
```python
# This cell is for grading. DO NOT remove it
```

## Reading a dataset from a text file (10 P)
In the previous task we created a method to classify the risk of patients, by manualy setting rules defining for which inputs the user is in `more` or `less` risk regarding their health. In the next exercises we will approach the task differently. Our goal is to create a classification method based on data. In order to achieve this we need to also create functions that loads the existing data into the program so that we can use it. Furthermore, we can use the loaded data to apply on our decision tree implementation and check what its outputs are.

The file `health-test.txt` contains several fictious records of personal data and habits. We split this task into two parts. In the first part, we assume that we have read a line from the file and can now process it. In the second function we load the file and process each line using the function we have defined for this purpose.

* Read the file automatically using the methods introduced during the lecture.
* Represent the dataset as a list of tuples. Make sure that the tuples have the same format as in the previous task, e.g. `('yes', 31, 'good')`.
* Make sure that you close the file after you have opened it and read its content. If you use a `with` statement then you don't have to worry about closing the file.

**Notes**: 
* Make sure when opening a file not to use an absolute path. An absolute path will
work on your computer, but when your code is tested on the departments computers it will fail. Use relative paths when opening files
* Values read from files are always strings.
* Each line contains a newline `\n` character at the end
* If you are using Windows as your operating system, refrain from opening any text files using Notepad. It will remove any linebreaks `\n`. You should inspect the files using the Jupyter text editor or any other modern text editor.

<!-- grade_id: cell-c1a8bc4c0e4ccb26, points: 0, grade: false, solution: true -->
```python
def parse_line_test(line):
    """
    Takes a line from the file, including a newline, and parses it into a patient tuple

    Args:
        line: A line from the `health-test.txt` file
    Returns:
        tuple: A tuple representing a patient
    """
    assert (
        line[-1] == "\n"
    ), "Did you change the contents of the line before calling this function?"
    # YOUR CODE HERE
    details = line.strip().split(',')

    details = [int(x) if x.isnumeric() else x for x in details]

    return tuple(details)

    # YOUR CODE HERE
    
```

<!-- grade_id: cell-4e1f7ad1e66b3121, points: 5, grade: true, solution: false -->
```python
x = "yes,23,good\n"
parsed_line = parse_line_test(x)
smoker, age, diet = parsed_line
print(parsed_line)
t.assertIsInstance(parsed_line, tuple)
t.assertEqual(len(parsed_line), 3)
t.assertIsInstance(age, int)
t.assertNotIn("\n", diet, "Are you handling line breaks correctly?")
t.assertEqual(parsed_line[-1], "good")

```

Output:
```
('yes', 23, 'good')
```

<!-- grade_id: cell-0ed341eaf90d76a6, points: 0, grade: true, solution: false -->
```python
# This cell is for grading. DO NOT remove it
```

<!-- grade_id: cell-844c38f934d4a3ef, points: 0, grade: false, solution: true -->
```python
def gettest():
    """
    Opens the `health-test.txt` file and parses it
    into a list of patient tuples. You are encouraged to use
    the `parse_line_test` function but it is not necessary to do so.

    This functions assumes that the `health-test.txt` file is located in
    the same directory as this notebook.

    Returns:
        A list of patient tuples as read from the file.
    """
    # YOUR CODE HERE

    list_of_patients = []

    f = open("health-test.txt", "r")
    for x in f:
        list_of_patients.append(parse_line_test(x))

    f.close()

    return list_of_patients
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-ab1b01fcaac26cd5, points: 5, grade: true, solution: false -->
```python
testset = gettest()
pprint(testset)
t.assertIsInstance(testset, list)
t.assertEqual(len(testset), 8)
t.assertIsInstance(testset[0], tuple)

```

Output:
```
[('yes', 21, 'poor'),
 ('no', 50, 'good'),
 ('no', 23, 'good'),
 ('yes', 45, 'poor'),
 ('yes', 51, 'good'),
 ('no', 60, 'good'),
 ('no', 15, 'poor'),
 ('no', 18, 'good')]
```

<!-- grade_id: cell-42b0a5cad68fc15a, points: 0, grade: true, solution: false -->
```python
# This cell is for grading. DO NOT remove it
```

## Applying the decision tree to the dataset (15 P)

* Apply the decision tree to all points in the dataset, and return the ratio of them that are classified as "more".
* A ratio is a value in [0-1]. So if out of 50 data points 15 return `"more"` the value that should be returned is `0.3`

<!-- grade_id: cell-6703ef98e2b5c93b, points: 0, grade: false, solution: true -->
```python
def evaluate_testset(dataset):
    """
    Calculates the percentage of datapoints for which the
    decision function evaluates to `more` for a given dataset

    Args:
        A list of patient tuples

    Returns:
        float: The percentage of data points which are evaluated to `'more'`
    """
    # YOUR CODE HERE
    count = 0
    for data in dataset:
        if decision(data) == 'more':
            count +=1
    
    return count/len(dataset)
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-c13a0b23c9faba52, points: 15, grade: true, solution: false -->
```python
ratio = evaluate_testset(gettest())
print(f"ratio --> {ratio}")
t.assertIsInstance(ratio, float)
assert_percentage(ratio)
t.assertTrue(0.3 < ratio < 0.4)

```

Output:
```
ratio --> 0.375
```

## Learning from examples (10 P)

Suppose that instead of relying on a fixed decision tree, we would like to use a data-driven approach where data points are classified based on a set of training observations manually labeled by experts. Such labeled dataset is available in the file `health-train.txt`. The first three columns have the same meaning than for `health-test.txt`, and the last column corresponds to the labels.

* Read the `health-train.txt` file and convert it into a list of pairs. The first element of each pair is a triplet of attributes (the patient tuple), and the second element is the label.
* Similarlly to the previous exercise we split the task into two parts. The first involves processing each line individually. The second handles opening the file and processing all lines of the file

**Note**: A triplet is a tuple that contains exactly three values, a pair is a tuple that contains exactly two values

<!-- grade_id: cell-fc38ed11fee6fbeb, points: 0, grade: false, solution: true -->
```python
def parse_line_train(line):
    """
    This function works similarly to the `parse_line_test` function.
    It parses a line of the `health-train.txt` file into a tuple that
    contains a patient tuple and a label.

    Args:
        line: A line from the `health-train.txt`

    Returns:
        A tuple that contains a patient tuple and a label as a string
    """
    assert line[-1] == "\n"
    # YOUR CODE HERE
    details = line.strip().split(',')

    details = [int(x) if x.isnumeric() else x for x in details]

    return (tuple(details[:3]),details[3])

    # YOUR CODE HERE
    
```

<!-- grade_id: cell-b97620858b167c0d, points: 5, grade: true, solution: false -->
```python
x = "yes,67,poor,more\n"
parsed_line = parse_line_train(x)
print(parsed_line)

t.assertIsInstance(parsed_line, tuple)
t.assertEqual(len(parsed_line), 2)

data, label = parsed_line

t.assertIsInstance(data, tuple)
t.assertEqual(len(data), 3)
t.assertEqual(data[1], 67)

t.assertIsInstance(label, str)
t.assertNotIn("\n", label, "Are you handling line breaks correctly?")
t.assertEqual(label, "more")

```

Output:
```
(('yes', 67, 'poor'), 'more')
```

<!-- grade_id: cell-ced9925599b5dd99, points: 0, grade: true, solution: false -->
```python
# This cell is for grading. DO NOT remove it
```

<!-- grade_id: cell-183f0bbeecd179c4, points: 0, grade: false, solution: true -->
```python
def gettrain():
    """
    Opens the `health-train.txt` file and parses it into
    a list of patient tuples accompanied by their respective label.

    Returns:
        A list of tuples comprised of a patient tuple and a label
    """
    # YOUR CODE HERE
    list_of_patients = []

    f = open("health-train.txt", "r")
    for x in f:
        list_of_patients.append(parse_line_train(x))

    f.close()

    return list_of_patients
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-a3d593f232e0403a, points: 5, grade: true, solution: false -->
```python
trainset = gettrain()
pprint(trainset)
t.assertIsInstance(trainset, list)
t.assertEqual(len(trainset), 16)
first_datapoint = trainset[0]
t.assertIsInstance(first_datapoint, tuple)
t.assertIsInstance(first_datapoint[0], tuple)
t.assertIsInstance(first_datapoint[1], str)
```

Output:
```
[(('yes', 54, 'good'), 'less'),
 (('no', 55, 'good'), 'less'),
 (('no', 26, 'good'), 'less'),
 (('yes', 40, 'good'), 'more'),
 (('yes', 25, 'poor'), 'less'),
 (('no', 13, 'poor'), 'more'),
 (('no', 15, 'good'), 'less'),
 (('no', 50, 'poor'), 'more'),
 (('yes', 33, 'good'), 'more'),
 (('no', 35, 'good'), 'less'),
 (('no', 41, 'good'), 'less'),
 (('yes', 30, 'poor'), 'more'),
 (('no', 39, 'poor'), 'more'),
 (('no', 20, 'good'), 'less'),
 (('yes', 18, 'poor'), 'less'),
 (('yes', 55, 'good'), 'more')]
```

<!-- grade_id: cell-7ab8ef2c931d69c1, points: 0, grade: true, solution: false -->
```python
# This cell is for grading. DO NOT remove it
```

## Nearest neighbor classifier (25 P)

We consider the nearest neighbor algorithm that classifies test points following the label of the nearest neighbor in the training data. You can read more about Nearest neighbor classifiers [here](http://www.robots.ox.ac.uk/~dclaus/digits/neighbour.htm). For this, we need to define a distance function between data points. We define it to be

`distance(a, b) = (a[0] != b[0]) + ((a[1] - b[1]) / 50.0) ** 2 + (a[2] != b[2])`

where `a` and `b` are two tuples representing two patients.

* Implement the distance function.
* Implement the function that retrieves for a test point the nearest neighbor in the training set, and classifies the test point accordingly (i.e. returns the label of the nearest data point).

**Hint**: You can use the special `infinity` floating point value with `float('inf')`

***Keep in mind that `bool`s in Python are also `int`s. `True` is the same as `1` and `False` is the same as `0`***

<!-- grade_id: cell-671ea24ec8a11241, points: 0, grade: false, solution: true -->
```python
def distance(a, b):
    """
    Calculates the distance between two data points (patient tuples)
    Args:
        a, b: Two patient tuples for which we want to calculate the distance
    Returns:
        float: The distance between a, b according to the above formula
    """
    # YOUR CODE HERE
    return (a[0] != b[0]) + ((a[1] - b[1]) / 50.0) ** 2 + (a[2] != b[2])
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-5929fce6ffc8ca0f, points: 5, grade: true, solution: false -->
```python
# Test distance
x1 = ("yes", 34, "poor")
x2 = ("yes", 51, "good")
dist = distance(x1, x2)
print(f"distance({x1}, {x2}) --> {dist}")
expected_dist = 1.1156
t.assertAlmostEqual(dist, expected_dist)

```

Output:
```
distance(('yes', 34, 'poor'), ('yes', 51, 'good')) --> 1.1156
```

<!-- grade_id: cell-8dc0c32bc78d1e62, points: 0, grade: true, solution: false -->
```python
# This cell is for grading. DO NOT remove it
```

<!-- grade_id: cell-e2ffaa5484ac1886, points: 0, grade: false, solution: true -->
```python
def neighbor(x, trainset):
    """
    Returns the label of the nearest data point in trainset to x.
    If x is `('no', 30, 'good')` and the nearest data point in trainset
    is `('no', 31, 'good')` with label `'less'` then `'less'` will be returned.
    In case two elements have the exact same distance, element that first occurs
    in the dataset is picked (the element with the smallest index).

    Args:
        x: The data point for which we want to find the nearest neighbor
        trainset: A list of tuples with patient tuples and a label

    Returns:
        str: The label of the nearest data point in the trainset.
        Can only be 'more' or 'less'
    """
    # YOUR CODE HERE
    distances = [distance(x,a[0]) for a in trainset]

    print(distances.index(min(distances)))

    return trainset[distances.index(min(distances))][1]
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-a36122337853f195, points: 10, grade: true, solution: false -->
```python
# Test neighbor
x = ("yes", 31, "good")
prediction = neighbor(x, gettrain())
print(f"prediction --> {prediction}")
expected = "more"
t.assertEqual(prediction, expected)

```

Output:
```
8
prediction --> more
```

<!-- grade_id: cell-cbf6ddfbfaab7875, points: 0, grade: true, solution: false -->
```python
# This cell is for grading. DO NOT remove it
```

In this part we want to compare the decision tree we have implemented with the nearest neighbor method. Apply both the decision tree and nearest neighbor classifiers on the test set, and return the list of data point(s) for which the two classifiers disagree, and with which probability it happens.

<!-- grade_id: cell-8dbf7da153f3d797, points: 0, grade: false, solution: true -->
```python
def compare_classifiers(trainset, testset):
    """
    This function compares the two classification methods (decision tree, nearest neighbor)
    by finding all the datapoints for which the methods disagree. It returns
    a list of the test datapoints for which the two methods do not return
    the same label as well as the ratio of those datapoints compared to the whole
    test set.

    Args:
        trainset: The training set used by the nearest neighbour classfier.
        testset: Contains the elements which will be used to compare the decision tree
        and nearest neighbor classification methods.

    Returns:
        A list containing all the data points which yield different results for the two
        classification methods. The ratio of datapoints for which the two methods
        disagree.

    """
    # YOUR CODE HERE
    dt_labels = [decision(x) for x in testset]
    nn_labels = [neighbor(x,trainset) for x in testset]

    difference = []

    for i, item in enumerate(testset):
        if(dt_labels[i] != nn_labels[i]):
            difference.append(item)

    return difference, len(difference)/len(testset)

    
    # YOUR CODE HERE
    
    return disagree, percentage
```

<!-- grade_id: cell-3b55f7e89ad4dfeb, points: 10, grade: true, solution: false -->
```python
# Test compare_classifiers
disagree, ratio = compare_classifiers(gettrain(), gettest())
print(f"ratio = {ratio}")
t.assertIsInstance(disagree, list)
t.assertIsInstance(ratio, float)
t.assertIsInstance(disagree[0], tuple)
t.assertEqual(len(disagree[0]), 3)
assert_percentage(ratio)
t.assertTrue(0.1 < ratio < 0.2)
```

Output:
```
14
1
2
11
0
1
5
13
ratio = 0.125
```

One problem of simple nearest neighbors is that one needs to compare the point to predict to all data points in the training set. This can be slow for datasets of thousands of points or more. Alternatively, some classifiers train a model first, and then use it to classify the data.

## Nearest mean classifier (25 P)

We consider one such trainable model, which operates in two steps:

1. Compute the average point for each class
2. Classify new points to be of the class whose average point is nearest to the point to predict.

For this classifier, we convert the attributes smoker and diet to real values (for smoker: `1.0` if 'yes' otherwise `0.0`, and for diet: `0.0` if 'good' otherwise `1.0`), and use the modified distance function:

`distance(a,b) = (a[0] - b[0]) ** 2 + ((a[1] - b[1]) / 50.0) ** 2 + (a[2] - b[2]) ** 2`

Age will also from now on be represented as a `float`. The new data points will be referred to as numerical patient tuples. 

We adopt an object-oriented approach for building this classifier.

* Implement the `gettrain_num` function that will load the training dataset from the `health-train.txt` file and parse each line to a numerical patient tuple with its label. You can still follow the same structure that we used before (i.e. using a `parse_line_...` function), however, it is not required for this exercise. Only the `gettrain_num` function will be tested.


* Implement the new distance function.


* Implement the methods `train` and `predict` of the class `NearestMeanClassifier`.

<!-- grade_id: cell-efadd1b300bd22ec, points: 0, grade: false, solution: true -->
```python
def parse_line_train_num(
    line,
):
    """
    Takes a line from the file `health-train.txt`, including a newline,
    and parses it into a numerical patient tuple.

    Args:
        line: A line from the `health-test.txt` file
    Returns:
        A numerical patient tuple and its label
    """
    # YOUR CODE HERE
    raise NotImplementedError("Replace this line with your code")
    # YOUR CODE HERE
    


def gettrain_num():
    """
    Parses the `health-train.txt` file into numerical patient tuples.

    Returns:
        A list of tuples containing numerical patient tuples and their labels
    """
    # YOUR CODE HERE
    raise NotImplementedError("Replace this line with your code")
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-b8a19b27d4c7de88, points: 5, grade: true, solution: false -->
```python
# Test gettrain_num
trainset_num = gettrain_num()
t.assertIsInstance(trainset_num, list)
first_datapoint = trainset_num[0]
print(f"first_datapoint --> {first_datapoint}")
t.assertIsInstance(first_datapoint[0], tuple)
t.assertIsInstance(first_datapoint[0][0], float)
t.assertIsInstance(first_datapoint[0][1], float)
t.assertIsInstance(first_datapoint[0][2], float)
```

<!-- grade_id: cell-789233fd9f800e77, points: 0, grade: true, solution: false -->
```python
# This cell is for grading. DO NOT remove it
```

<!-- grade_id: cell-d35765538d512a1a, points: 0, grade: false, solution: true -->
```python
def distance_num(a, b):
    """
    Calculates the distance between two numerical patient tuples.
    Args:
        a, b: Two numerical patient tuples for which we want to calculate the distance.
    Returns:
        float: The distance between a, b.
    """
    # YOUR CODE HERE
    raise NotImplementedError("Replace this line with your code")
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-de4d8b77145f3bff, points: 2, grade: true, solution: false -->
```python
x1 = (1.0, 23.0, 0.0)
x2 = (0.0, 41.0, 1.0)
dist = distance_num(x1, x2)
print(f"dist --> {dist}")
t.assertIsInstance(dist, float)
t.assertTrue(2.12 < dist < 2.13)
```

<!-- grade_id: cell-28628e7373da18f2, points: 0, grade: true, solution: false -->
```python
# This cell is for grading. DO NOT remove it
```

<!-- grade_id: cell-e0b339bfd0fcc16c, points: 0, grade: false, solution: true -->
```python
class NearestMeanClassifier:
    """
    Represents a NearestMeanClassifier.

    When an instance is trained a dataset is provided and the mean for each class is calculated.
    During prediction the instance compares the datapoint to each class mean (not all datapoints)
    and returns the label of the class mean to which the datapoint is closest to.

    Instance Attributes:
        more: A tuple representing the mean of
        every 'more' data-point in the dataset
        less: A tuple representing the mean of
        every 'less' data-point in the dataset
    """

    def __init__(self):
        self.more = None
        self.less = None

    def train(self, dataset):
        """
        Calculates the class means for a given dataset and stores
        them in instance attributes more, less.

        The mean of the more class is a tuple containing three elements.
        Each element of the mean tuple contains the mean of all the elements
        in the training set that are labeled `more` for each corresponding index.
        This means that the mean tuple contains the mean smoker, age and health
        values.
        The same is true of the less mean tuple, but for all the elements
        labeled `less`.

        This function is does not return anything useful, but it has the side
        effect of setting the more and less instance variables.

        Args:
            dataset: A list of tuples each of them containing a numerical patient tuple
            and its label
        Returns:
            self
        """
        # YOUR CODE HERE
        raise NotImplementedError("Replace this line with your code")
        # YOUR CODE HERE
        
        return self

    def predict(self, x):
        """
        Returns a prediction/label for numeric patient tuple x.
        The classifier compares the given data point to the mean
        class tuples of each class and returns the label of the
        class to which x is the closest to (according to our
        distance function).

        Args:
            x: A numerical patient tuplefor which we want a prediction
        Returns:
            str: The predicted label.
        """
        # YOUR CODE HERE
        raise NotImplementedError("Replace this line with your code")
        # YOUR CODE HERE
        

    def __repr__(self):
        try:
            more = tuple(round(m, 3) for m in self.more)
            less = tuple(round(l, 3) for l in self.less)
        except AttributeError:
            more, less = None, None
        return f"{self.__class__.__name__}(more: {more}, less: {less})"
```

* Instantiate the `NearestMeanClassifier`, train it on the training data, and return it

<!-- grade_id: cell-5f7f00ee83c94703, points: 0, grade: false, solution: true -->
```python
def build_and_train(trainset_num):
    """
    Instantiates the `NearestMeanClassifier`, trains it on the
    `trainset_num` dataset and returns it.

    Args:
        trainset_num: A list of numerical patient tuples with their respective labels

    Returns:
        NearestMeanClassifier: A NearestMeanClassifier trained on `trainset_num`
    """
    # YOUR CODE HERE
    raise NotImplementedError("Replace this line with your code")
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-415891bde4cbde19, points: 2, grade: true, solution: false -->
```python
# Test build_and_train
classifier = NearestMeanClassifier()
classifier.train(trainset_num)
t.assertIsInstance(classifier, NearestMeanClassifier)
print(classifier)
try:
    classifier.more, classifier.less
except AttributeError:
    t.fail(
        "Did you train the classifier?"
        " Did you store the mean vector for the 'more' class?"
        " Did you store the mean vector for the 'less' class?",
    )

t.assertIsInstance(classifier.more, tuple)
t.assertIsInstance(classifier.less, tuple)

t.assertEqual(round(classifier.more[1]), 37)
t.assertEqual(round(classifier.less[1]), 32)

```

<!-- grade_id: cell-ca2d0921e96ede25, points: 6, grade: true, solution: false -->
```python
# This cell is for grading. Do NOT remove it
```

* Load the test dataset into memory as a list of numerical patient tuples
* Predict the test data using the nearest mean classifier and return the index of all test examples for which all three classifiers (decision tree, nearest neighbor and nearest mean) agree.

**Note**: Be careful that the `NearestMeanClassifier` expects the dataset in a different form, compared to the other two methods.

<!-- grade_id: cell-f37f3035a32a8f85, points: 0, grade: false, solution: true -->
```python
def gettest_num():
    """
    Parses the `health-test.txt` file into numerical patient tuples.

    Returns:
        list: A list containing numerical patient tuples, loaded from `health-test.txt`
    """

    # YOUR CODE HERE
    raise NotImplementedError("Replace this line with your code")
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-f3656461c994dc3d, points: 2, grade: true, solution: false -->
```python
testset_num = gettest_num()
pprint(testset_num)
t.assertIsInstance(testset_num, list)
t.assertEqual(len(testset_num), 8)
t.assertIsInstance(testset_num[0], tuple)
t.assertEqual(len(testset_num[0]), 3)
```

<!-- grade_id: cell-008fcc0d21f07c3a, points: 0, grade: false, solution: true -->
```python
def predict_test():
    """
    Classifies the test set using all the methods that were developed in this exercise sheet,
    namely `decision`, `neighbor` and `NearestMeanClassifier`.

    This functions loads all the needed data by calling the corresponding functions.
    (gettrain, gettest, gettrain_num, gettest_num)

    Returns:
        A list of the indices of all the datapoints for which all three classfiers have
        the same output.
    """
    # YOUR CODE HERE
    raise NotImplementedError("Replace this line with your code")
    # YOUR CODE HERE
    
    return agreed_samples
```

<!-- grade_id: cell-853c957eaaf81c28, points: 8, grade: true, solution: false -->
```python
same_predictions = predict_test()
pprint(same_predictions)
t.assertIsInstance(same_predictions, list)
t.assertEqual(len(same_predictions), 6)
t.assertIsInstance(same_predictions[0], int)
t.assertEqual(same_predictions[-2:], [6, 7])
```